**Important notes:**

- Activate the conda environment `mr` before you start: `conda activate mr`

- Don't change the name of the file and don't delete any cells.

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and (if necessary) collaborators below.

- The function `NotImplementedError()` prevents you from hand in assignments with empty cells. Simply delete the function if you start working on a cell with this entry.

- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells:
  - if you use a *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`
  - if you use *Visual Studio Code*: select "Restart" and then "Run All" 

Good luck!

```python
NAME = ""
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Netflix user engagement analyse

## Setup

```python
# Wir importieren die Bibliothek pandas
import pandas as pd
```

---
## Daten

### Daten importieren (2 Punkte)

- Wir haben Pandas mit der Bezeichnung `pd` importiert.

- Die Pandas-Funktion zum Einlesen von CSV-Dateien: `.read_csv`

- CSV-Daten mit Pandas importieren und mit der Bezeichnung `df` speichern. 

- Pfad zu den Daten in GitHub: 'https://raw.githubusercontent.com/kirenz/datasets/master/netflix.csv'


- Hinweis:

```python
___ = ___.___(___)
```

<!-- grade_id: import, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: correct_import, points: 2, grade: true, solution: false -->
```python
"""Check that df returns the correct output"""
assert len(df) == 2000
assert df.columns.tolist() == ['group', 'outcome']
```

---
## Analyse 

### Kreuztabelle (4 Punkte)

Schreiben Sie den Code, der eine Kreuztabelle erzeugt erzeugt:

- Die Funktion in Pandas zur Erzeugung von Kreuztabellen: `.crosstab`
- Speichern Sie die Tabelle unter der Bezeichnung: `cross_table` 
- Zeilen: `group`
- Spalten: `outcome` 
- Geben Sie die Spalten- und Zeilensummen aus (`margins`) 
- Geben Sie die Werte als Prozentwerte aus (`normalize`)
- Die Variable `group` soll in der Tabelle "Gruppe" genannt werden (`rownames`)
- Die Variable `group` soll in der Tabelle "Gruppe" genannt werden (`colnames`)
- Formatierung der Prozent-Werte: bspw. sollte ein Wert von 0.1775 als 17.75 angezeigt werden
- Geben Sie die Tabelle aus

- Hinweis:

```python

___ = __.__(__['___'], ___['___'],
            ___=True,
            ___=True,
            ___=['___'],  # Umbenennung der Zeilen
            ___=['___'] # Umbenennung der Spalten
            ) * ___ # Multiplikation mit 100 für eine bessere Darstellung der Prozentwerte

___

```

<!-- grade_id: cross_table, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: correct_cross_table, points: 4, grade: true, solution: false -->
```python
"""Check that it returns the correct output"""
assert cross_table['Engagement'].B == 21.8
```